---
title       : Chapter 1
description : Insert the chapter description here
attachments :
  slides_link : http://s3.amazonaws.com/assets.datacamp.com/production/course_3555/slides/chapter1.pdf



--- type:VideoExercise lang:r xp:50 skills:1 key:72d9b55c9c
## Welcome to the thing

Some stuff here

*** =projector_key
92d179019a6d5934563c11dcf7fa6370

--- type:NormalExercise lang:r xp:100 skills:1 key:ce4af30374
## Exercise 1
  
*** =instructions
- These
- are
- some better
- instructions

*** =hint
- add hint
*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}
print('test')
```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:f97e6945bb
## Exercise 2


*** =instructions
- One
- Two
- Three
- Four

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:789590b254
## Exercise 3

*** =projector_key
2a3c287ae67a5daca58bef719782ef70

--- type:VideoExercise lang:r xp:50 skills:1 key:097ac2f591
## This is a video

*** =projector_key
5fca8b90ce32115ff40b2a4a4755c121
