---
title       : Chapter 1
description : Insert the chapter description here
attachments :
  slides_link : http://s3.amazonaws.com/assets.datacamp.com/production/course_3555/slides/chapter1.pdf



--- type:VideoExercise lang:r xp:50 skills:1 key:72d9b55c9c
## Welcome to the thing

Some stuff here

*** =slides_key
92d179019a6d5934563c11dcf7fa6370

--- type:NormalExercise lang:r xp:100 skills:1 key:ce4af30374
## Exercise 1

<script>alert('test')</script>

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:f97e6945bb
## Exercise 2


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:789590b254
## Exercise 3

*** =slides_key
2a3c287ae67a5daca58bef719782ef70
