---
title       : Chapter 2
description : Insert the chapter description here

--- type:NormalExercise lang:r xp:100 skills:1 key:72d9b55c9c
## Exercise 1


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:f97e6945bb
## Exercise 2

This is an exercise

*** =instructions
- One
- Two
- Three
- Four

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:789590b254
## Exercise 3

